# Activity: Variables in Python

## Introduction

In this activity you will practice the following:
- Assigning specific values / types to variables
- Overwriting variables

#### Question 1

Create variable `name` and assign any value you want to it.

<!-- grade_id: cell-89bbfa8c8528e855, points: 0, grade: false, solution: true -->
```python
# Your code here
name =('green')
print(name)
```

Output:
```
green
```

<!-- grade_id: cell-3e323501e904d29b, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

print('Make sure you created the variable: name')
print(name)


```

Output:
```
Make sure you created the variable: name
NameError                                 Traceback (most recent call last)
Cell In[3], line 4
      1 # Question 1 Grading Checks
      3 print('Make sure you created the variable: name')
----> 4 print(name)

NameError: name 'name' is not defined
```

#### Question 2

Assign the value of `2.7` to the variable `seconds`.

<!-- grade_id: cell-8d5cb3e863fba294, points: 0, grade: false, solution: true -->
```python
# Your code here
seconds = ('2.7')
print (seconds)
```

Output:
```
2.7
```

<!-- grade_id: cell-ea1c3340532802f6, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert isinstance(seconds, float), 'Have you assigned a float value to the variable seconds?'

```

#### Question 3

Assign boolean value `True` to the variable `python_is_fun`.

<!-- grade_id: cell-513f546ddd5fd8a9, points: 0, grade: false, solution: true -->
```python
python_is_fun = True
print ('python_is_fun')
```

Output:
```
python_is_fun
```

<!-- grade_id: cell-5ed340946941a67f, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert isinstance(python_is_fun, bool), 'Have you assigned a boolean value to the python_is_fun variable?'
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[5], line 3
      1 # Question 3 Grading Checks
----> 3 assert isinstance(python_is_fun, bool), 'Have you assigned a boolean value to the python_is_fun variable?'

NameError: name 'python_is_fun' is not defined
```